---
title: 'Word vectors: mapping words to numerical vectors '
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to word vectors

```yaml
type: VideoExercise
key: 8390294b9a
xp: 50
```

`@projector_key`
e493fe72b09a643d6cc7f8e7b6c447bb

---

## Skip-gram and CBOW algorithms

```yaml
type: VideoExercise
key: 5b0f2c66fd
xp: 50
```

`@projector_key`
02ca9345b662923d873a71500884bea9
